# SLU06 - Dealing with Data Problems

This notebook has exercises covering the following topics:

- Tidy Data
- Data Entry Problems
- Missing Values

```python
import os
import pandas as pd
import numpy as np
import hashlib
import json
```

The dataset that we'll use for these exercises comes from the World Health Organisation, and records the counts of confirmed tuberculosis cases by country, year (between 1980 and 2008), and demographic group.
The demographic groups are broken down by sex (m, f) and age (0–14, 15–25, 25–34, 35–44, 45–54, 55–64, 65+, unknown).

But as you will see, this dataset doesn't follow the Tidy Data Principle.

In the following exercises, we will clean it.

First let's read the dataset into a pandas dataframe.

```python
df = pd.read_csv(os.path.join('data', 'tb.csv'))
df.head(10)
```

Output:
```
country  year  new_sp_m014  new_sp_m1524  new_sp_m2534  new_sp_m3544  \
0      AD  1989          NaN           NaN           NaN           NaN   
1      AD  1990          NaN           NaN           NaN           NaN   
2      AD  1991          NaN           NaN           NaN           NaN   
3      AD  1992          NaN           NaN           NaN           NaN   
4      AD  1993          NaN           NaN           NaN           NaN   
5      AD  1994          NaN           NaN           NaN           NaN   
6      AD  1996          0.0           0.0           0.0           4.0   
7      AD  1997          0.0           0.0           1.0           2.0   
8      AD  1998          0.0           0.0           0.0           1.0   
9      AD  1999          0.0           0.0           0.0           1.0   

   new_sp_m4554  new_sp_m5564  new_sp_m65  new_sp_mu  new_sp_f014  \
0           NaN           NaN         NaN        NaN          NaN   
1           NaN           NaN         NaN        
```

### Exercise 1

Create a list with the column names that refer to the demographic groups data (i.e, those that start with `new_sp_`).

Call the list `demo_group_data`. The values in list should be in the same order as they are in the dataframe's columns.

```python
demo_group_data = list(df[[i for i in df.columns if "new_sp_" in i]])


```

<!-- grade_id: cell-188948d13a5b357c, points: 0, grade: false, solution: true -->
```python
# Create a list with the demographic groups data
# demo_group_data = ...

# YOUR CODE HERE


demo_group_data = list(df[[col for col in df.columns if "new_sp_" in col]])



```

<!-- grade_id: cell-d828d2d6cb20d296, points: 2, grade: true, solution: false -->
```python
first_element_hash = 'c8ae9368386d7c77204db0840e77d432aff39941d6f42407b0164028b0bdd5c6'
entire_list_hash = '11db8cc5591bcd9757e052b44e7b1a5ccd4f2de2e5c3130d5437834ea021aa78'

assert isinstance(demo_group_data, list), "demo_group_data should be a list."

assert len(demo_group_data) == 16, "demo_group_data doesn't have the right number of elements."

error_msg = 'The first element of the list is not correct.'
assert first_element_hash == hashlib.sha256(bytes(demo_group_data[0], encoding='utf8')).hexdigest(), error_msg

error_msg = 'The list is not correct.'
assert entire_list_hash == hashlib.sha256(json.dumps(demo_group_data).encode()).hexdigest()
```

### Exercise 2

Create a new dataframe, that is the result of changing `df` so that it has 4 columns:
- country
- year
- demo_group
- cases

The values in the demo_group column should be exactly the ones that we currently have as column names.

The new dataframe should be called `df_tidy_1`.

Hint: use the answer from **Exercise 1**.

```python
df
```

Output:
```
country  year  new_sp_m014  new_sp_m1524  new_sp_m2534  new_sp_m3544  \
0         AD  1989          NaN           NaN           NaN           NaN   
1         AD  1990          NaN           NaN           NaN           NaN   
2         AD  1991          NaN           NaN           NaN           NaN   
3         AD  1992          NaN           NaN           NaN           NaN   
4         AD  1993          NaN           NaN           NaN           NaN   
5         AD  1994          NaN           NaN           NaN           NaN   
6         AD  1996          0.0           0.0           0.0           4.0   
7         AD  1997          0.0           0.0           1.0           2.0   
8         AD  1998          0.0           0.0           0.0           1.0   
9         AD  1999          0.0           0.0           0.0           1.0   
10        AD  2000          0.0           0.0           1.0           0.0   
11        AD  2001          0.0           NaN           NaN           2.0   
12  
```

<!-- grade_id: cell-3e7fa8d07fe433cc, points: 0, grade: false, solution: true -->
```python
# Create a new version of df that doesn't have variable values as columns
# df_tidy_1 = ...

# YOUR CODE HERE

df_tidy_1 = pd.melt(df, id_vars= ['country','year'],  value_vars =  demo_group_data, var_name = 'demo_group', value_name= 'cases')
```

```python
print(df_tidy_1)
```

Output:
```
country  year   demo_group  cases
0          AD  1989  new_sp_m014    NaN
1          AD  1990  new_sp_m014    NaN
2          AD  1991  new_sp_m014    NaN
3          AD  1992  new_sp_m014    NaN
4          AD  1993  new_sp_m014    NaN
5          AD  1994  new_sp_m014    NaN
6          AD  1996  new_sp_m014    0.0
7          AD  1997  new_sp_m014    0.0
8          AD  1998  new_sp_m014    0.0
9          AD  1999  new_sp_m014    0.0
10         AD  2000  new_sp_m014    0.0
11         AD  2001  new_sp_m014    0.0
12         AD  2002  new_sp_m014    0.0
13         AD  2003  new_sp_m014    0.0
14         AD  2004  new_sp_m014    0.0
15         AD  2005  new_sp_m014    0.0
16         AD  2006  new_sp_m014    0.0
17         AD  2007  new_sp_m014    NaN
18         AD  2008  new_sp_m014    0.0
19         AE  1980  new_sp_m014    NaN
20         AE  1981  new_sp_m014    NaN
21         AE  1982  new_sp_m014    NaN
22         AE  1983  new_sp_m014    NaN
23         AE  1984  new_sp_m014    NaN
24    
```

<!-- grade_id: cell-640d2b3b55c2cc39, points: 2, grade: true, solution: false -->
```python
column_names_hash = '8a5277b30342779427d2bf91bdbe9669975eecaf85c3560103367a75af3d971c'
dataframe_shape_hash = 'd28d50e42c0b9b03762856921319b41c42f7fa4d1b5da5c5b4b84d716eff3a9c'
dataframe_hash = '1c564c7ad9e350444637e70de9af00877e99dc5518dc130664617e4ce51b1a42'

assert isinstance(df_tidy_1, pd.DataFrame), "df_tidy_1 should be a pandas DataFrame."

error_msg = "The dataframe doesn't have the right columns."
assert column_names_hash == hashlib.sha256(json.dumps(df_tidy_1.columns.tolist()).encode()).hexdigest(), error_msg

error_msg = "The dataframe doesn't have the right shape."
assert dataframe_shape_hash == hashlib.sha256(json.dumps(df_tidy_1.shape).encode()).hexdigest(), error_msg

error_msg = "The dataframe is not correct."
assert dataframe_hash == hashlib.sha256(df_tidy_1.to_msgpack()).hexdigest(), error_msg
```

### Exercise 3

Now, take a look at the values in column `demo_group`.

They all start with a meaningless "new_sp_" string, so we will remove it.

Create a new dataframe `df_tidy_2`, that is a copy of `df_tidy_1`, but where the `demo_group` column no longer has the "new_sp_" in each value. 

```python
df_tidy_1['demo_group'] = df_tidy_1['demo_group'].str.replace('new_sp_', '')

```

<!-- grade_id: cell-9a0680701ef7a9e8, points: 0, grade: false, solution: true -->
```python
# Create a new version of df_tidy_1 that doesn't have "new_sp_" in the demo_group column
# df_tidy_2 = ...

# YOUR CODE HERE
df_tidy_1['demo_group'] = df_tidy_1['demo_group'].str.replace('new_sp_', '')
df_tidy_2 = pd.DataFrame(df_tidy_1)
```

<!-- grade_id: cell-b63a0ff46979395c, points: 2, grade: true, solution: false -->
```python
column_names_hash = '8a5277b30342779427d2bf91bdbe9669975eecaf85c3560103367a75af3d971c'
dataframe_shape_hash = 'd28d50e42c0b9b03762856921319b41c42f7fa4d1b5da5c5b4b84d716eff3a9c'
dataframe_hash = '5d7508a0b5cda0e5f4cb530f5922199576d255ac7340d9fc1c89b28b5d6b355d'

assert isinstance(df_tidy_2, pd.DataFrame), "df_tidy_2 should be a pandas DataFrame."

error_msg = "The dataframe doesn't have the right columns."
assert column_names_hash == hashlib.sha256(json.dumps(df_tidy_2.columns.tolist()).encode()).hexdigest(), error_msg

error_msg = "The dataframe doesn't have the right shape."
assert dataframe_shape_hash == hashlib.sha256(json.dumps(df_tidy_2.shape).encode()).hexdigest(), error_msg

error_msg = "The dataframe is not correct."
assert dataframe_hash == hashlib.sha256(df_tidy_2.to_msgpack()).hexdigest(), error_msg
```

### Exercise 4

As you may have noticed, our dataset still has a problem. The `demo_group` column has data that in fact represents two variables: `gender` and `age`.

So our end goal will be to replace the `demo_group` column with the two new columns.

On this exercise, create a new dataframe `df_tidy_3`, that is a copy of `df_tidy_2`, but has a new column `gender`, which has the first letter of column `demo_group` (m/f).

Hint: you may need to search for this one online, as you'll need a string method that we didn't cover in the Learning notebook.

<!-- grade_id: cell-c9c176a8e3461c10, points: 0, grade: false, solution: true -->
```python
# Create a new version of df_tidy_2 that has a new column gender
# df_tidy_3 = ...

# YOUR CODE HERE

df_tidy_2['gender'] = df_tidy_2['demo_group'].astype(str).str[0]
df_tidy_3 = pd.DataFrame(df_tidy_2)
```

<!-- grade_id: cell-c06912eb8d895afc, points: 2, grade: true, solution: false -->
```python
column_names_hash = 'a3efb4e4eb3a43fa947a95032395ccc21d2cc3715a8d359e4c711132f4b837e0'
dataframe_shape_hash = '78cdfddff871199ede18988405d0daf7c840512ff0bdb01168d3f85394618ef8'
dataframe_hash = '282c81e0a303c5515386d4d78bbd2307377447e9f14ea35d6bcbdc2ec720b1bb'

assert isinstance(df_tidy_3, pd.DataFrame), "df_tidy_3 should be a pandas DataFrame."

error_msg = "The dataframe doesn't have the right columns."
assert column_names_hash == hashlib.sha256(json.dumps(df_tidy_3.columns.tolist()).encode()).hexdigest(), error_msg

error_msg = "The dataframe doesn't have the right shape."
assert dataframe_shape_hash == hashlib.sha256(json.dumps(df_tidy_3.shape).encode()).hexdigest(), error_msg

error_msg = "The dataframe is not correct."
assert dataframe_hash == hashlib.sha256(df_tidy_3.to_msgpack()).hexdigest(), error_msg
```

### Exercise 5

Now we want to create the column `age` and get rid of column `demo_group`.

On this exercise, create a new dataframe `df_tidy_4`, that is a copy of `df_tidy_3`, but has a new column `age`, which is the same as column `demo_group`, except that it has the first letter removed.

We also want to get rid of column `demo_group`, so make sure the new dataframe doesn't have it!

Hint: you may need to search for this one online, as you'll need a string method that we didn't cover in the Learning notebook.

<!-- grade_id: cell-0a1b561a9c60b159, points: 0, grade: false, solution: true -->
```python
# Create a new version of df_tidy_3 that has a new column age, and doesn't have column demo_group
# df_tidy_4 = ...

# YOUR CODE HERE
df_tidy_3['age'] = df_tidy_2['demo_group'].astype(str).str[1:]
df_tidy_3 = df_tidy_3.drop(['demo_group'], axis = 1)
df_tidy_4 = pd.DataFrame(df_tidy_3)

```

<!-- grade_id: cell-52d4f44d3697b15a, points: 2, grade: true, solution: false -->
```python
column_names_hash = 'ab05d50deb2a842b806a607ad26166c82f91960bd9dae704d8a11ff78fb60860'
dataframe_shape_hash = '78cdfddff871199ede18988405d0daf7c840512ff0bdb01168d3f85394618ef8'
dataframe_hash = 'f534bc81fdd39b7a79c0630d6a3a0728b89c40d6c95bf297d94e537bfa305f84'

assert isinstance(df_tidy_4, pd.DataFrame), "df_tidy_4 should be a pandas DataFrame."

error_msg = "The dataframe doesn't have the right columns."
assert column_names_hash == hashlib.sha256(json.dumps(df_tidy_4.columns.tolist()).encode()).hexdigest(), error_msg

error_msg = "The dataframe doesn't have the right shape."
assert dataframe_shape_hash == hashlib.sha256(json.dumps(df_tidy_4.shape).encode()).hexdigest(), error_msg

error_msg = "The dataframe is not correct."
assert dataframe_hash == hashlib.sha256(df_tidy_4.to_msgpack()).hexdigest(), error_msg
```

### Exercise 6

If you take a look at the age values, they are quite hard to understand... 
Let's make those easier to grasp!

Create a new dataframe `df_tidy_5`, where the `age` column has more understandable values.

We want that each age interval is separated by a `-`. For instance, `014` should be replaced with `0-14`.
We also want the last age group to be represented as `65+`. The unknown values `u` can be left unchanged.

```python
df_tidy_4.age.value_counts()
```

Output:
```
1524    11538
u       11538
014     11538
2534    11538
65      11538
5564    11538
4554    11538
3544    11538
Name: age, dtype: int64
```

<!-- grade_id: cell-0a5e28ea0faa0992, points: 0, grade: false, solution: true -->
```python
# Create a new version of df_tidy_4 where the age values are more understandable
# df_tidy_5 = ...

# YOUR CODE HERE
ca = []
for i in df_tidy_4['age']:
    if i == 'u':
        ca.append(i)
    elif i == '65':
        ca.append(i + '+')
    else:
        ca.append(str(i[:len(i)-2]) + '-' + i[-2:])

df_tidy_4['age'] = ca
df_tidy_5 = pd.DataFrame(df_tidy_4)      
        
print(df_tidy_5)

```

Output:
```
country  year  cases gender   age
0          AD  1989    NaN      m  0-14
1          AD  1990    NaN      m  0-14
2          AD  1991    NaN      m  0-14
3          AD  1992    NaN      m  0-14
4          AD  1993    NaN      m  0-14
5          AD  1994    NaN      m  0-14
6          AD  1996    0.0      m  0-14
7          AD  1997    0.0      m  0-14
8          AD  1998    0.0      m  0-14
9          AD  1999    0.0      m  0-14
10         AD  2000    0.0      m  0-14
11         AD  2001    0.0      m  0-14
12         AD  2002    0.0      m  0-14
13         AD  2003    0.0      m  0-14
14         AD  2004    0.0      m  0-14
15         AD  2005    0.0      m  0-14
16         AD  2006    0.0      m  0-14
17         AD  2007    NaN      m  0-14
18         AD  2008    0.0      m  0-14
19         AE  1980    NaN      m  0-14
20         AE  1981    NaN      m  0-14
21         AE  1982    NaN      m  0-14
22         AE  1983    NaN      m  0-14
23         AE  1984    NaN      m  0-14
24    
```

<!-- grade_id: cell-794006afade645c7, points: 2, grade: true, solution: false -->
```python
column_names_hash = 'ab05d50deb2a842b806a607ad26166c82f91960bd9dae704d8a11ff78fb60860'
dataframe_shape_hash = '78cdfddff871199ede18988405d0daf7c840512ff0bdb01168d3f85394618ef8'
dataframe_hash = '324f52dadcd96598cd0edd64ce2ea6717aa7263a105a4b53b7ee5f6ea8baee30'

assert isinstance(df_tidy_4, pd.DataFrame), "df_tidy_5 should be a pandas DataFrame."

error_msg = "The dataframe doesn't have the right columns."
assert column_names_hash == hashlib.sha256(json.dumps(df_tidy_5.columns.tolist()).encode()).hexdigest(), error_msg

error_msg = "The dataframe doesn't have the right shape."
assert dataframe_shape_hash == hashlib.sha256(json.dumps(df_tidy_5.shape).encode()).hexdigest(), error_msg

error_msg = "The dataframe is not correct."
assert dataframe_hash == hashlib.sha256(df_tidy_5.to_msgpack()).hexdigest(), error_msg
```

### Exercise 7

Now that our dataset follows the Tidy Data principles, let's check for duplicates.

Assign to variable `duplicates_count` the number of duplicates in `df_tidy_5`. Make sure the value you get is an integer.

<!-- grade_id: cell-77ab90e009cd292d, points: 0, grade: false, solution: true -->
```python
# Count the number of duplicates in df_tidy_5
# duplicates_count = ...
# YOUR CODE HERE
duplicates = df_tidy_5.duplicated(keep='first')
duplicates_count = int(duplicates.sum())


```

<!-- grade_id: cell-b559fc483c9c8be1, points: 2, grade: true, solution: false -->
```python
duplicates_count_hash = 'e3b0c44298fc1c149afbf4c8996fb92427ae41e4649b934ca495991b7852b855'

error_msg = "duplicates_count must be an integer (explicitly convert it to 'int' if it is type 'numpy.int64')"
assert isinstance(duplicates_count, int), error_msg

error_msg = "duplicates_count doesn't have the right value."
assert duplicates_count_hash == hashlib.sha256(bytes(duplicates_count)).hexdigest(), error_msg
```

### Exercise 8

Now let's count the number of rows with missing values in our `df_tidy_5` dataframe.

Assign that value to variable `null_count` and make sure it is an integer.

Also take a moment to think about the number of null values we had in this dataset, in comparison to the dataset size.

<!-- grade_id: cell-526e2589fa58add5, points: 0, grade: false, solution: true -->
```python
# Count the number of null values in df_tidy_5
# null_count = ...

# YOUR CODE HERE
n_count_row = df_tidy_5.isnull().any(axis=1)
null_count = int(n_count_row.values.sum())
```

<!-- grade_id: cell-3da7144cdc77c7b9, points: 2, grade: true, solution: false -->
```python
duplicates_count_hash = 'e41482b02884d9b42b5ea9abb63be3658c28266e8f62f31db33fc9efa53cf01d'

error_msg = "null_count must be an integer (explicitly convert it to 'int' if it is type 'numpy.int64')"
assert isinstance(null_count, int), error_msg

error_msg = "null_count doesn't have the right value."
assert duplicates_count_hash == hashlib.sha256(bytes(null_count)).hexdigest(), error_msg
```

### Exercise 9

Now let's use a common technique to imput missing values: replacing them with the mean.

Create a new dataframe `df_tidy_6`, which is a copy of `df_tidy_5`, except that the missing values in the `cases` column are replaced with the mean of the column.

Also, take a moment to think if this is a good strategy in this case.

```python
df_tidy_5['cases'] = df_tidy_5.cases.fillna(df_tidy_5.cases.mean())
df_tidy_6 = pd.DataFrame(df_tidy_5)
print(df_tidy_6)
```

Output:
```
country  year       cases gender   age
0          AD  1989  636.764374      m  0-14
1          AD  1990  636.764374      m  0-14
2          AD  1991  636.764374      m  0-14
3          AD  1992  636.764374      m  0-14
4          AD  1993  636.764374      m  0-14
5          AD  1994  636.764374      m  0-14
6          AD  1996    0.000000      m  0-14
7          AD  1997    0.000000      m  0-14
8          AD  1998    0.000000      m  0-14
9          AD  1999    0.000000      m  0-14
10         AD  2000    0.000000      m  0-14
11         AD  2001    0.000000      m  0-14
12         AD  2002    0.000000      m  0-14
13         AD  2003    0.000000      m  0-14
14         AD  2004    0.000000      m  0-14
15         AD  2005    0.000000      m  0-14
16         AD  2006    0.000000      m  0-14
17         AD  2007  636.764374      m  0-14
18         AD  2008    0.000000      m  0-14
19         AE  1980  636.764374      m  0-14
20         AE  1981  636.764374      m  0-14
21         AE  1
```

<!-- grade_id: cell-91a85cda9755c1f5, points: 0, grade: false, solution: true -->
```python
# Create a new version of df_tidy_5 where the cases missing values are replaced with the mean
# df_tidy_6 = ...

# YOUR CODE HERE

df_tidy_5['cases'] = df_tidy_5.cases.fillna(df_tidy_5.cases.mean())
df_tidy_6 = pd.DataFrame(df_tidy_5)

```

<!-- grade_id: cell-5f3a449dd2f19c7d, points: 2, grade: true, solution: false -->
```python
column_names_hash = 'ab05d50deb2a842b806a607ad26166c82f91960bd9dae704d8a11ff78fb60860'
dataframe_shape_hash = '78cdfddff871199ede18988405d0daf7c840512ff0bdb01168d3f85394618ef8'
dataframe_hash = '1f8e04b67a127979a1a47c9aac44716d9c08b9aa0d7641b7bb8542769a32db37'

assert isinstance(df_tidy_6, pd.DataFrame), "df_tidy_6 should be a pandas DataFrame."

error_msg = "The dataframe doesn't have the right columns."
assert column_names_hash == hashlib.sha256(json.dumps(df_tidy_6.columns.tolist()).encode()).hexdigest(), error_msg

error_msg = "The dataframe doesn't have the right shape."
assert dataframe_shape_hash == hashlib.sha256(json.dumps(df_tidy_6.shape).encode()).hexdigest(), error_msg

error_msg = "The dataframe is not correct."
assert dataframe_hash == hashlib.sha256(df_tidy_6.to_msgpack()).hexdigest(), error_msg
```

### Exercise 10

Did you think that replacing the missing values with the mean was a terrible idea in this case? Because it is!

Let's see why.

Create a new dataframe `top_cases`, by:
- sorting `df_tidy_5` by the cases column in descending order
- taking the first row per country, using the drop_duplicates function
- taking the first 10 rows of that dataframe

Then take a look at the countries in `top_countries`. Can you see a correlation between the number of cases and the population size of those countries?

<!-- grade_id: cell-21d03667d23b1548, points: 0, grade: false, solution: true -->
```python
# Create a dataframe with the rows that correspond to the 10 highest number of cases in df_tidy_5
# top_cases = ...

# YOUR CODE HERE


top_cases = pd.DataFrame()
top_cases = top_cases.head(10)
```

Output:
```
<ipython-input-60-1e9b09956dfc> in <module>
      5 
      6 
----> 7 top_cases = pd.DataFrame(df_tidy_5, df_tidy_5.sort(by='cases',ascending = False))
      8 top_cases = top_cases.country.drop_duplicated(keep='first').shape
      9 top_cases = top_cases.head(10)

/anaconda3/envs/slu06/lib/python3.7/site-packages/pandas/core/generic.py in __getattr__(self, name)
   5065             if self._info_axis._can_hold_identifiers_and_holds_name(name):
   5066                 return self[name]
-> 5067             return object.__getattribute__(self, name)
   5068 
   5069     def __setattr__(self, name, value):

AttributeError: 'DataFrame' object has no attribute 'sort'
```

<!-- grade_id: cell-042e21c24cf691b9, points: 2, grade: true, solution: false -->
```python
column_names_hash = 'ab05d50deb2a842b806a607ad26166c82f91960bd9dae704d8a11ff78fb60860'
dataframe_shape_hash = '36d0cc685f3df22cfee85b856d878e734000d5cf3ff3cc3f809a96c623410ff4'
dataframe_hash = '8420df6d826597ab69f60b0c1131a440908a1fbe7ed5abdb673c17040ad72f4a'

assert isinstance(top_cases, pd.DataFrame), "top_cases should be a pandas DataFrame."

error_msg = "The dataframe doesn't have the right columns."
assert column_names_hash == hashlib.sha256(json.dumps(top_cases.columns.tolist()).encode()).hexdigest(), error_msg

error_msg = "The dataframe doesn't have the right shape."
assert dataframe_shape_hash == hashlib.sha256(json.dumps(top_cases.shape).encode()).hexdigest(), error_msg

error_msg = "The dataframe is not correct."
assert dataframe_hash == hashlib.sha256(top_cases.to_msgpack()).hexdigest(), error_msg
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-61-35e81e403372> in <module>
      9 
     10 error_msg = "The dataframe doesn't have the right shape."
---> 11 assert dataframe_shape_hash == hashlib.sha256(json.dumps(top_cases.shape).encode()).hexdigest(), error_msg
     12 
     13 error_msg = "The dataframe is not correct."

AssertionError: The dataframe doesn't have the right shape.
```

### Exercise 11 - ungraded

Take some time to discuss with a coleague sitting next to you, what would be a better way to impute missing values in this dataset.

```python

```